# Tutorial 2: RDB

**The goal of this tutorial is to implement 5 relation algebra operators and create 5 queries with them.**

The database is stored as a pickle file, a serialization protocol for Python data structures:

You can find the structure of the database in the example section below.

https://docs.python.org/3/library/pickle.html#examples

These are the relation algebra operators:

* __Set Union (A ∪ B)__: combine the tuple of both A and B
* __Set Difference (A - B)__: keep tuples in A that are not in B
* __Set Intersection (A ∩ B)__: find tuples present both in A and B
* __Selection (R σ COND)__: filter tuples in R that satisfies COND
* __Projection (R π COLS)__: select attributes in R mentioned in COLS
* __Join/Product (R X S)__: generate all combinations of R and S tuples

You can think of relational algebra as pseudo-code for database queries.

These operators are used in every database systems, under different names.

I encourage you to consider the complexity of each operator to find optimal queries !

__Grade Scale__: 20 points
* correct operator/query: 2 point
* incorrect operator/query: 0 point

__Further documentations__:

* https://www.imdb.com/interfaces/
* https://learnxinyminutes.com/docs/python/
* https://en.wikipedia.org/wiki/Relational_algebra
* https://docs.python.org/3/tutorial/datastructures.html
* https://www.dataquest.io/blog/jupyter-notebook-tutorial/

# Core

```python
# import pickle from standard library
import pickle

# load the database as a pickle file
with open('imdb.pickle', 'rb') as r:
    DB = pickle.load(r)
    
# limit the amount of results to 20
def limit(q):  # decorator
    def f(*args, **kwargs):
        return q(*args, **kwargs)[:20]
    return f
```

# Examples

```python
# the database is a dictionnary
# print the names of the tables
list(DB.keys())
```

Output:
```
['names', 'basics', 'akas', 'ratings', 'writers', 'directors', 'principals']
```

```python
# a table is a list of tuples
# retrieve the first 3 tuples
DB["names"][0:3]
```

Output:
```
[('nm0000007',
  'Humphrey Bogart',
  1899,
  1957,
  'actor,soundtrack,producer',
  'tt0037382,tt0033870,tt0034583,tt0043265'),
 ('nm0000026',
  'Cary Grant',
  1904,
  1986,
  'actor,soundtrack,producer',
  'tt0053125,tt0036613,tt0048728,tt0056923'),
 ('nm0000033',
  'Alfred Hitchcock',
  1899,
  1980,
  'actor,director,producer',
  'tt0054215,tt0053125,tt0052357,tt0030341')]
```

```python
# print the first table row
for name, table in DB.items():
    for row in table:
        print(name)
        print(row)
        break
```

Output:
```
names
('nm0000007', 'Humphrey Bogart', 1899, 1957, 'actor,soundtrack,producer', 'tt0037382,tt0033870,tt0034583,tt0043265')
basics
('tt0100275', 'movie', 'The Wandering Soap Opera', 'La Telenovela Errante', 0, 2017, None, 80, 'Comedy,Drama,Fantasy')
akas
('tt0100275', 1, 'La Telenovela Errante', None, None, 'original', None, 1)
ratings
('tt0100275', 6.8, 92)
writers
('tt0100275', 'nm0749914')
directors
('tt0100275', 'nm0749914')
principals
('tt0100275', 1, 'nm0016013', 'actor', None, None)
```

```python
# there is one important thing you should know about Python values !
# most data structures are mutable, and can be modified in your code

a = [1, 2, 3]
a.extend([4, 5, 6])

a
```

Output:
```
[1, 2, 3, 4, 5, 6]
```

```python
# you can prevent these side effects by creating:
# - an empty structure
# - a copy

a = [1, 2, 3]
b = a.copy()
c = list()

b.extend([4, 5, 6])
c.extend([4, 5, 6])

a, b, c
```

Output:
```
([1, 2, 3], [1, 2, 3, 4, 5, 6], [4, 5, 6])
```

```python
# you can also use immutable operators to create new object

a = [1, 2, 3]
b = a + [4, 5, 6]

a, b
```

Output:
```
([1, 2, 3], [1, 2, 3, 4, 5, 6])
```

# OPERATORS

```python
# this operator is provided as an example

def union(r, s):
    """Concatenate tuples from relation r and s."""
    t = list()
    t.extend(r)
    t.extend(s)
    return t
```

```python
assert union([], []) == []
assert union([(1, 2, 3)], []) == [(1, 2, 3)]
assert union([], [(1, 2, 3)]) == [(1, 2, 3)]
assert union([(1, 2, 3)], [(4, 5, 6)]) == [(1, 2, 3), (4, 5, 6)]
assert union([(4, 5, 6)], [(1, 2, 3)]) == [(4, 5, 6), (1, 2, 3)]
assert union([(1, 2, 3), (4, 5, 6)], [(7, 8, 9)]) == [(1, 2, 3), (4, 5, 6), (7, 8, 9)]
assert union([(1, 2, 3)], [(4, 5, 6), (7, 8, 9)]) == [(1, 2, 3), (4, 5, 6), (7, 8, 9)]
```

```python
assert union([1,2,3],[1,2,3]) == [1,2,3]
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-22-d124f7184495> in <module>()
----> 1 assert union([1,2,3],[1,2,3]) == [1,2,3]

AssertionError:
```

<!-- grade_id: cell-80e16a9e439489ca, points: 0, grade: false, solution: true -->
```python
def difference(r, s):
    """Keep the tuple of r that are not in s."""
    t = list()
    t.extend(r)
    # YOUR CODE HERE
    i=0
    
    while i<len(t):
        loopbroken = False
        for j in range(len(s)):
            if t[i] == s[j]:
                t.remove(t[i])
                loopbroken = True
                break
        if loopbroken == False:
            i+=1
            
    #raise NotImplementedError()
    return t
```

<!-- grade_id: cell-f03c5807a027b112, points: 2, grade: true, solution: false -->
```python
assert difference([], []) == []
assert difference([(1, 2, 3)], []) == [(1, 2, 3)]
assert difference([], [(1, 2, 3)]) == []
assert difference([(1, 2, 3)], [(1, 2, 3)]) == []
assert difference([(1, 2, 3)], [(4, 5, 6)]) == [(1, 2, 3)]
assert difference([(1, 2, 3)], [(1, 2, 3), (4, 5, 6)]) == []
assert difference([(1, 2, 3)], [(4, 5, 6), (7, 8, 8)]) == [(1, 2, 3)]
assert difference([(1, 2, 3), (4, 5, 6)], [(1, 2, 3)]) == [(4, 5, 6)]
assert difference([(1, 2, 3), (4, 5, 6)], [(4, 5, 6)]) == [(1, 2, 3)]
assert difference([(1, 2, 3), (4, 5, 6)], [(7, 8, 9)]) == [(1, 2, 3), (4, 5, 6)]
```

<!-- grade_id: cell-a66a7ba43d252280, points: 0, grade: false, solution: true -->
```python
def intersection(r, s):
    """Keep tuples of r that are also in s."""
    t = list()
    # YOUR CODE HERE
    for i in range(len(s)):
        for j in range(len(r)):
            if s[i] == r[j]:
                t.append(s[i])
            
    #raise NotImplementedError()
    return t
```

<!-- grade_id: cell-a8b09737ac98c2df, points: 2, grade: true, solution: false -->
```python
assert intersection([], []) == []
assert intersection([(1, 2, 3)], []) == []
assert intersection([], [(1, 2, 3)]) == []
assert intersection([(1, 2, 3)], [(4, 5, 6)]) == []
assert intersection([(4, 5, 6)], [(1, 2, 3)]) == []
assert intersection([(1, 2, 3)], [(1, 2, 3)]) == [(1, 2, 3)]
assert intersection([(4, 5, 6)], [(4, 5, 6)]) == [(4, 5, 6)]
assert intersection([(1, 2, 3), (4, 5, 6)], [(7, 8, 9)]) == []
assert intersection([(1, 2, 3), (4, 5, 6)], [(1, 2, 3)]) == [(1, 2, 3)]
assert intersection([(1, 2, 3), (4, 5, 6)], [(4, 5, 6)]) == [(4, 5, 6)]
```

<!-- grade_id: cell-c4df43a7b0d7c4e8, points: 0, grade: false, solution: true -->
```python
def selection(r, cond):
    """Keep tuples that satisfy cond (i.e., when cond is True)."""
    t = list()
    # YOUR CODE HERE
    for i in range(len(r)):
        if cond(r[i]):
            t.append(r[i])
    #raise NotImplementedError()
    return t
```

<!-- grade_id: cell-988b583e38ada6d8, points: 2, grade: true, solution: false -->
```python
assert selection([], lambda x: True) == []
assert selection([], lambda x: False) == []
assert selection([(1, 2, 3)], lambda x: False) == []
assert selection([(4, 5, 6)], lambda x: False) == []
assert selection([(1, 2, 3)], lambda x: x[0] == 9) == []
assert selection([(1, 2, 3)], lambda x: True) == [(1, 2, 3)]
assert selection([(4, 5, 6)], lambda x: True) == [(4, 5, 6)]
assert selection([(1, 2, 3)], lambda x: x[0] == 1) == [(1, 2, 3)]
assert selection([(1, 2, 3), (4, 5, 6)], lambda x: x[0] == 9) == []
assert selection([(1, 2, 3), (4, 5, 6)], lambda x: x[0] == 1) == [(1, 2, 3)]
assert selection([(1, 2, 3), (4, 5, 6)], lambda x: x[0] == 4) == [(4, 5, 6)]
```

<!-- grade_id: cell-08d77510c95184b4, points: 0, grade: false, solution: true -->
```python
def projection(r, cols):
    """Keep attributes that are in cols for each tuple of r."""
    t = list()
    # YOUR CODE HERE
    for i in range(len(r)):
        x=[]
        for j in range(len(cols)):
            x.append(r[i][cols[j]])
        t.append(tuple(x))
            
    #raise NotImplementedError()
    return t
```

<!-- grade_id: cell-ca83359ee07e113b, points: 2, grade: true, solution: false -->
```python
assert projection([], []) == []
assert projection([(4, 5, 6)], []) == [()]
assert projection([(4, 5, 6)], [0]) == [(4,)]
assert projection([(4, 5, 6)], [1]) == [(5,)]
assert projection([(4, 5, 6)], [2]) == [(6,)]
assert projection([(4, 5, 6)], [0, 2]) == [(4, 6)]
assert projection([(4, 5, 6), (7, 8, 9)], [1]) == [(5,), (8,)]
assert projection([(4, 5, 6), (7, 8, 9)], [0, 2]) == [(4, 6), (7, 9)]
```

<!-- grade_id: cell-82b890bfc576cef7, points: 0, grade: false, solution: true -->
```python
def join(r, s):
    """Generate all combination of tuples for relation r and s."""
    t = list()
    # YOUR CODE HERE
    for i in range(len(r)):
        for j in range(len(s)):
            temp = list(r[i])
            temp.extend(list(s[j]))
            t.append(tuple(temp))
            #t.append(tuple(list(r[i]).extend(list(s[j]))))
    #raise NotImplementedError()
    return t
```

<!-- grade_id: cell-5395f5033f5793db, points: 2, grade: true, solution: false -->
```python
assert join([], []) == []
assert join([(1, 2, 3)], []) == []
assert join([], [(1, 2, 3)]) == []
assert join([(1, 2, 3)], [(4, 5, 6)]) == [(1, 2, 3, 4, 5, 6)]
assert join([(4, 5, 6)], [(1, 2, 3)]) == [(4, 5, 6, 1, 2, 3)]
assert join([(1, 2, 3), (4, 5, 6)], [(7, 8, 9)]) == [(1, 2, 3, 7, 8, 9), (4, 5, 6, 7, 8, 9)]
assert join([(7, 8, 9)], [(1, 2, 3), (4, 5, 6)]) == [(7, 8, 9, 1, 2, 3), (7, 8, 9, 4, 5, 6)]
assert join([(1, 2), (3, 4)], [(5, 6), (7, 9), (9, 0)]) == [(1, 2, 5, 6), (1, 2, 7, 9),
                                                            (1, 2, 9, 0), (3, 4, 5, 6),
                                                            (3, 4, 7, 9), (3, 4, 9, 0)]
```

# QUERIES

__0. Select the primary title, start year and runtime of movies that are 120 minutes long__
  - _hint: always filter a query with `selection` before applying any other operators_

```python
@limit
def Q0():
    r1 = selection(DB["basics"], lambda x: x[7] is not None and x[7] == 120)
    r2 = projection(r1, [2, 5, 7])
    return r2

Q0()
```

Output:
```
[('Justice League', 2017, 120),
 ('Five Fingers for Marseilles', 2017, 120),
 ('Long Live the Horror', 2017, 120),
 ('Wonderkind', 2017, 120),
 ('Snowflake', 2017, 120),
 ('The Dinner', 2017, 120),
 ('Patel Ki Punjabi Shaadi', 2017, 120),
 ('Beyond the Clouds', 2017, 120),
 ('Myr vashomu domu!', 2017, 120),
 ('The Man with the Iron Heart', 2017, 120),
 ('G: A Dark Tale of Desires', 2017, 120),
 ('Bank Chor', 2017, 120),
 ('Vorticale', 2017, 120),
 ('Love Pret-a-porte', 2017, 120),
 ('Blood Ties', 2017, 120),
 ('Mary Shelley', 2017, 120),
 ('The Space Between Us', 2017, 120),
 ('Okja', 2017, 120),
 ("Fate/Stay Night: Heaven's Feel - I. Presage Flower", 2017, 120),
 ('Rough Stuff', 2017, 120)]
```

```python
assert Q0() == [
    ('Justice League', 2017, 120),
    ('Five Fingers for Marseilles', 2017, 120),
    ('Long Live the Horror', 2017, 120),
    ('Wonderkind', 2017, 120),
    ('Snowflake', 2017, 120),
    ('The Dinner', 2017, 120),
    ('Patel Ki Punjabi Shaadi', 2017, 120),
    ('Beyond the Clouds', 2017, 120),
    ('Myr vashomu domu!', 2017, 120),
    ('The Man with the Iron Heart', 2017, 120),
    ('G: A Dark Tale of Desires', 2017, 120),
    ('Bank Chor', 2017, 120),
    ('Vorticale', 2017, 120),
    ('Love Pret-a-porte', 2017, 120),
    ('Blood Ties', 2017, 120),
    ('Mary Shelley', 2017, 120),
    ('The Space Between Us', 2017, 120),
    ('Okja', 2017, 120),
    ("Fate/Stay Night: Heaven's Feel - I. Presage Flower", 2017, 120),
    ('Rough Stuff', 2017, 120),
]
```

__1: Select the name of persons born in 2000 and whose primary profession is 'actresses'__
  - _hint: you can query the 'names' table to retrieve these two information_

<!-- grade_id: cell-473734dbb058859c, points: 0, grade: false, solution: true -->
```python
@limit
def Q1():
    # YOUR CODE HERE
    r1 = selection(DB["names"], lambda x: x[2] == 2000 and 'actress' == x[4]) #in?
    r2 = projection(r1, [1])
    #raise NotImplementedError()
    return r2

Q1()
```

Output:
```
[('Esme Creed-Miles',),
 ('Cami Ottman',),
 ('Shelby Lyon',),
 ('Minami Hamabe',),
 ('Moka Kamishiraishi',),
 ('Bente Fokkens',),
 ('Mima Ito',),
 ('Na-Na OuYang',),
 ('Zaira Wasim',),
 ('Destina Baser',)]
```

<!-- grade_id: cell-dab152dc1fee5624, points: 2, grade: true, solution: false -->
```python
assert len(Q1()) == 10  # has 10 rows
assert all(len(row) == 1 for row in Q1())  # has 1 columns per row (primary name)
```

__2: Select the name, rating and votes of movies whose rating > 9 and number of vote > 1000__
  - _hint: you should join the `ratings` and `basics` relations to associated their information_

<!-- grade_id: cell-1877b3fd18660e79, points: 0, grade: false, solution: true -->
```python
@limit
def Q2():
    # YOUR CODE HERE
    r1 = selection(DB["ratings"],lambda x: x[1]>9 and x[2]>1000)
    r2 = projection(DB["basics"],[0,2])
    r3 = join (r2, r1)
    r3 = selection(r3, lambda x: x[0] == x[2])
    
    return projection(r3,[1,3,4])
    #raise NotImplementedError()

Q2()
```

Output:
```
[('Hans Zimmer: Live in Prague', 9.1, 1293),
 ('Aloko Udapadi', 9.6, 6435),
 ('On vam ne Dimon', 9.2, 2618)]
```

<!-- grade_id: cell-e36b40d053907d08, points: 2, grade: true, solution: false -->
```python
assert len(Q2()) == 3  # has 3 rows
assert all(len(row) == 3 for row in Q2())  # has 3 column per row (primary title, average rating, number of votes)
```

__3: Select the primary name and genre of movies directed by 'Larry Rosen'__
  - _hint: remember to filter tuple before joining relations together !_

<!-- grade_id: cell-71ed082b62c037ff, points: 0, grade: false, solution: true -->
```python
@limit
def Q3():
    # YOUR CODE HERE
    r1 = selection(DB["names"], lambda x: x[1] == 'Larry Rosen')
    r1 = projection(r1,[0,1])
    #raise NotImplementedError()
    r2 = join(DB["directors"],r1)
    r2 = selection(r2, lambda x: x[1] == x[2])
    
    r3 = projection(DB["basics"], [0,8,2])
    r3 = join(r3,r2)
    r3 = selection(r3, lambda x: x[0] == x[3])
    
    r4 = projection(r3, [2,1])
    

    return r4
Q3()
```

Output:
```
[('After the Outbreak', 'Horror,Sci-Fi'),
 ('Paranoia Tapes', 'Horror'),
 ('Surviving the Outbreak', 'Horror,Sci-Fi'),
 ('The New Roommate', 'Drama,Thriller'),
 ('Into the Outbreak', 'Horror,Sci-Fi'),
 ('Paranoia Films 2: Press Play', 'Horror'),
 ('Death at a Barbecue', 'Horror,Thriller'),
 ('Second Escape', 'Drama'),
 ('Something Like Love', 'Drama'),
 ('Gwendolyn', 'Drama'),
 ('Revenge is Best Served', 'Horror'),
 ('The Question', 'Drama,Romance'),
 ('Death of Love', 'Drama,Romance')]
```

<!-- grade_id: cell-4bea5e25f426a53a, points: 2, grade: true, solution: false -->
```python
assert len(Q3()) == 13  # has 13 rows
assert all(len(row) == 2 for row in Q3())  # has 2 columns per row (primary title, genres)
```

__4: Select the translated title and language of the movie: 'Minecraft the Christmas Movie'__
  - _hint: you can find these information in the `akas` table_

<!-- grade_id: cell-1ff71a1dcc3378bf, points: 0, grade: false, solution: true -->
```python
@limit
def Q4():
    # YOUR CODE HERE
    #raise NotImplementedError()
    r1 = selection(DB["basics"], lambda x: x[2] == 'Minecraft the Christmas Movie')
    r1 = projection(r1, [0,2])
    
    r2 = projection(DB["akas"], [0,2,4])
    
    r3 = join(r2,r1)
    r3 = selection(r3, lambda x: x[0] == x[3])
    r3 = projection(r3, [1,2])
    
    return r3

Q4()
```

Output:
```
[('Minecraft la película de navidad', None),
 ('Minecraft the Christmas Movie', None),
 ("Minecraft Rozhdestvenskiy fil'm", None),
 ('Minecraft the Christmas Movie', None),
 ('Mainkurafutokurisumasumubi', None),
 ('Minecraft Filmul de Craciun', None),
 ('Minecraft the movie', None),
 ('Minecraft Le film de Noël', 'fr'),
 ('Minecraft Der Weihnachtsfilm', None),
 ('Minecraft I tainía ton Christougénnon', None)]
```

<!-- grade_id: cell-61f0b5a2b0edf17d, points: 2, grade: true, solution: false -->
```python
assert len(Q4()) == 10  # has 10 rows
assert all(len(row) for row in Q4())  # has 2 rows (title, language)
```

__5: Select the primary title of movies with a rating > 9.5 but no translations (akas)__
  - _hint: this is a job for set operators !_

<!-- grade_id: cell-193f9d1070c87e8a, points: 0, grade: false, solution: true -->
```python
@limit
def Q5():
    r1 = selection(DB["ratings"], lambda x : x[1]>9.5)
    r2 = projection(DB["akas"],[0,2])
    
    r4 = projection(DB["basics"],[0,2])
    r4 = join(r4,r1)
    r4 = selection(r4, lambda x: x[0] == x[2])
    
    return projection(r4,[1])

Q5()
```

Output:
```
[('The Musician',),
 ('The Narcissists',),
 ('Aloko Udapadi',),
 ('Moda Mia',),
 ('Lina',),
 ('Ufrivillig',),
 ('Golani',),
 ('To Let',),
 ('MacDonald Ranch',),
 ('Pupa',),
 ('Taawdo the Sunlight',),
 ('Una Historia Necesaria',),
 ('Re-action',),
 ('Trobocop: H synomwsia tou petradiou',),
 ('Ego-Sum',),
 ('Never-Ending Road',),
 ('Gangter in Morteni',)]
```

<!-- grade_id: cell-79928838dd6c2845, points: 2, grade: true, solution: false -->
```python
assert len(Q5()) == 17
assert all(len(row) == 1 for row in Q5())  # has 1 columns per row (primary title)
```