# Tutorial 3: NoSQL

__The goal of this assignment is to create 10 "queries" based on 2 NoSQL databases.__

I put queries in quotes, as these databases do not provide a declarative query language such as SQL.

Instead, you must rely on the database API to extract information and combine the results with some Python code.

During this tutorial, I encourage you to consider the pros and cons of each databases compared to relational databases.

__The first database: imdb_basics.shelve contains the movies records of IMDB (basics)__.

This database is stored using `shelve`, a module of Python standard library.

`shelve` is a simple Key-Value store that acts as a persistent dictionary.

This is the same model as:
- Redis
- Memcached
- Google LevelDB
- Amazon DynamoDB
- Facebook RocksDB

__The second database: (imdb.json) contains the person records of IMDB (names)__.

This database is stored using `tinymongo`, a drop in replacement for MongoDB.

`tinymongo` is a Document store that provides query methods on JSON docs.

This is the same model as:
- MongoDB
- CouchDB
- ArangoDB
- RethinkDB
- ElasticSearch

__These two data models are [the most popular ones](https://db-engines.com/en/ranking) and they are also used to create more [advanced data model](https://github.com/datathings/greycat/tree/master/plugins)__.

__Grade scale__: 20 points
- correct query: 2 point
- incorrect query: 0 points

__Further documentations__:
* https://www.imdb.com/interfaces/
* https://learnxinyminutes.com/docs/python/
* https://github.com/schapman1974/tinymongo/
* https://docs.python.org/3.6/library/shelve.html

# Core

```python
# import shelve from standard library
import shelve

# open a connection to shelve database
basics = shelve.open('imdb_basics.shelve', 'r')
```

```python
# import shelve from external dependency
from tinymongo import TinyMongoClient

# open a connection to tinymongo database
names = TinyMongoClient('.').imdb.names
```

# Examples

```python
# get the first 20 keys from basics collection
list(basics.keys())[:20]
```

Output:
```
['tt0100275',
 'tt0112502',
 'tt0137204',
 'tt0331314',
 'tt0339736',
 'tt0366415',
 'tt0368133',
 'tt0451201',
 'tt0451279',
 'tt0460890',
 'tt0491175',
 'tt0491203',
 'tt0493405',
 'tt0498381',
 'tt0835802',
 'tt0874268',
 'tt0937225',
 'tt0974015',
 'tt1020581',
 'tt1024855']
```

```python
# get the first item from the basics collection
basics['tt0100275']
```

Output:
```
{'_id': 'tt0100275',
 'titleType': 'movie',
 'primaryTitle': 'The Wandering Soap Opera',
 'originalTitle': 'La Telenovela Errante',
 'isAdult': 0,
 'startYear': 2017,
 'endYear': None,
 'runtimeMinutes': 80,
 'genres': ['Comedy', 'Drama', 'Fantasy']}
```

```python
# get the first item from the names collection
names.find_one()
```

Output:
```
{'_id': 'nm7691973',
 'primaryName': 'Baste Granfon',
 'birthYear': 2012,
 'deathYear': None,
 'knownForTitles': ['tt7736492', 'tt5577086', 'tt6766834', 'tt8207358'],
 'primaryProfession': ['actor']}
```

# Queries

__1. How many movies are in the `basics` collection ?__
- _hint_: you don't have to use a loop to answer this query
- _return_: Count (where Count = int)

<!-- grade_id: cell-ed73e2be3e703d89, points: 0, grade: false, solution: true -->
```python
def Q1():
    # YOUR CODE HERE
    # raise NotImplementedError()
    
    return len(basics.keys())

Q1()
```

Output:
```
18199
```

<!-- grade_id: cell-610479179ac3aa45, points: 2, grade: true, solution: false -->
```python
assert isinstance(Q1(), int)
```

__2. Select the record associated to the movie whose primaryTitle is 'Blade Runner 2049'__
* _hints_: you have to use method the database provides
* _return_: Record (where Record = Dict[str, Any])

<!-- grade_id: cell-e58d8a8565370bfa, points: 0, grade: false, solution: true -->
```python
def Q2():
    # YOUR CODE HERE
    # raise NotImplementedError()
    for i in basics.keys():
        if basics[i]['primaryTitle'] == 'Blade Runner 2049':
            return basics[i]

Q2()
```

Output:
```
{'_id': 'tt1856101',
 'titleType': 'movie',
 'primaryTitle': 'Blade Runner 2049',
 'originalTitle': 'Blade Runner 2049',
 'isAdult': 0,
 'startYear': 2017,
 'endYear': None,
 'runtimeMinutes': 164,
 'genres': ['Drama', 'Mystery', 'Sci-Fi']}
```

<!-- grade_id: cell-b8841a1fe6ff032f, points: 2, grade: true, solution: false -->
```python
assert isinstance(Q2(), dict)
assert Q2()['primaryTitle'] == 'Blade Runner 2049'
```

__3. Select the primary title and runtime of every movies longer than 300 minutes (excluded)__
* _hint_: you have to construct your own return value
* _return_: List[Tuple[primaryTitle, runtimeMinutes]] (where primaryTitle = str, runtimeMinutes=int)

<!-- grade_id: cell-824c1acfafdceedf, points: 0, grade: false, solution: true -->
```python
def Q3():
    # YOUR CODE HERE
    # raise NotImplementedError()
    out = []
    for i in basics.keys():
        if int(basics[i]['runtimeMinutes'] or 0) > 300:
            out.append((basics[i]['primaryTitle'],basics[i]['runtimeMinutes']))
        
    return out
Q3()
```

Output:
```
[('Who was Hitler', 450),
 ('Painting', 360),
 ('Next Stop', 406),
 ('Mazwara', 788),
 ('h36:', 2160),
 ('Stuck', 360),
 ('Nari', 6017),
 ('Europa: The Last Battle', 746),
 ('Sakhi', 1179),
 ('An Infants Journey: Reggio Emilia Approach', 912),
 ('1998: The Deadliest Year For Children In American History', 396),
 ('Laundry', 359),
 ('Make Me Fly', 623),
 ('Bullfighting Memories', 1100)]
```

<!-- grade_id: cell-92156cd5f8035a2e, points: 2, grade: true, solution: false -->
```python
assert len(Q3()) == 14
assert all(len(row) == 2 and row[1] > 300 for row in Q3())
```

__4. Select the record in the `names` collection associated to \_id 'nm0705356'__
* _hint_: use `find_one` to return only one record
* _return_: Record (where Record = Dict[str, Any])

<!-- grade_id: cell-56a09933c73fe5ff, points: 0, grade: false, solution: true -->
```python
def Q4():
    # YOUR CODE HERE
    #raise NotImplementedError()
    return names.find_one({"_id":"nm0705356"})

Q4()
```

Output:
```
{'_id': 'nm0705356',
 'primaryName': 'Daniel Radcliffe',
 'birthYear': 1989,
 'deathYear': None,
 'knownForTitles': ['tt0373889', 'tt4034354', 'tt0926084', 'tt1201607'],
 'primaryProfession': ['actor', 'soundtrack', 'producer']}
```

<!-- grade_id: cell-de600da118a6dec7, points: 2, grade: true, solution: false -->
```python
assert isinstance(Q4(), dict)
assert Q4()['_id'] == 'nm0705356'
```

__5. Select the primaryName of the first 20 persons born in 2000, sorted by name (descending)__
* _hint_: use the `find` method to return multiple results
* _return_: List[Name] (where Name = str)

<!-- grade_id: cell-e667fccbb5d8a574, points: 0, grade: false, solution: true -->
```python
def Q5():
    # YOUR CODE HERE
    # raise NotImplementedError()
    out = []
    for i in names.find({'birthYear': 2000}):
        out.append(i['primaryName'])
    out.sort(reverse = True)
    return out[:20]

Q5()
```

Output:
```
['Zaira Wasim',
 'Willow Shields',
 'Viktor Derek',
 'Tobias Nikolai Haugland',
 'Stanislaw Cywka',
 'Shelby Lyon',
 'Ryan Henry Knight',
 'Rachelle Henry',
 'Pedro Diego',
 'Na-Na OuYang',
 'Moka Kamishiraishi',
 'Minami Hamabe',
 'Mima Ito',
 'Max Baissette de Malglaive',
 'Margaret Manousos',
 'Ludvig Fahlstedt',
 'Joshua Latorro',
 'Jonah Bryson',
 'Jackson Yee',
 'G.M. Whiting']
```

<!-- grade_id: cell-038b270abaff25ce, points: 2, grade: true, solution: false -->
```python
assert len(Q5()) == 20
assert all(isinstance(x, str) for x in Q5())
```

__6. Select the primaryName and birthYear of persons born after 2000 (excluded) and whose name starts with the letter 'M'__
* _hint_: use the `$and`, `$gt` and `$regex` operator of MongoDB
* _return_: List[Tuple[primaryName, birthYear]] (where primaryName = str, birthYear = int)

<!-- grade_id: cell-b9bf55def2db660b, points: 0, grade: false, solution: true -->
```python
def Q6():
    # YOUR CODE HERE
    # raise NotImplementedError()
    out = []
    for i in names.find({"$and":[{'birthYear':{"$gt":2000} },{'primaryName':{"$regex":'M.'}}]}):
        out.append((i['primaryName'],i['birthYear']))
    return out

Q6()
```

Output:
```
[('Mckenna Grace', 2006),
 ('Moxie Jillette', 2005),
 ('Maddie Dixon-Poirier', 2005),
 ('Mana Ashida', 2004),
 ('Mujtuba Ahmed', 2004),
 ('Mariangeli Collado', 2003),
 ('Madison Wolfe', 2002),
 ('Maisa Silva', 2002),
 ('Milo Parker', 2002)]
```

<!-- grade_id: cell-1cea9289869013f0, points: 2, grade: true, solution: false -->
```python
assert len(Q6()) == 9
assert all(len(row) == 2 for row in Q6())
assert all(row[1] > 2000 for row in Q6())
```

__7. Compute the average movie runtime in minutes__
* _hint_: aggregation has to be performed with code
* _return_: Average (where Average = float)

<!-- grade_id: cell-2b3d8c908851350f, points: 0, grade: false, solution: true -->
```python
def Q7():
    # YOUR CODE HERE
    # raise NotImplementedError()
    suma = 0.0
    for i in basics.keys():
        suma += float(basics[i]['runtimeMinutes'] or 0.0)
    return suma/len(basics.keys())

Q7()
```

Output:
```
61.72877630639046
```

<!-- grade_id: cell-b18c8bf9308ba4ce, points: 2, grade: true, solution: false -->
```python
assert isinstance(Q7(), float)
```

__8. Select the primary name and the primary titles for which the first 20 persons are known for__
* _hint_: you have to join the two database collections
* _return_: List[Tuple[primaryName, List[primaryTitle]]] (where primaryName = primaryTitle = str)

<!-- grade_id: cell-12e87ffb9860b577, points: 0, grade: false, solution: true -->
```python
def Q8():
    # YOUR CODE HERE
    # raise NotImplementedError()
    out = []
    for i in names.find():
        temp = i['knownForTitles']
        temp_l = []
        for j in temp:
            try:
                temp_l.append(basics[j]['primaryTitle'])
            except KeyError:
                pass
        out.append((i['primaryName'],temp_l))
        
    return out[:20]

Q8()
```

Output:
```
[('Baste Granfon', ['Meant to Beh']),
 ('Nicole Abbey', ['Bakers Club']),
 ('Valentina Gordon', ['Bluest Moon']),
 ('Viktor Antonio', ['Spunti na vode']),
 ('Chutian Liu', []),
 ('Run-yin Bai',
  ['The Bold, the Corrupt, and the Beautiful',
   'Mr. Long',
   'A Fish Out of Water',
   'See You After Life']),
 ('Aybars Kartal Özson', ['Benim Babam Bir Melek']),
 ('Fiorela Duranda', []),
 ('Robert Levey II', ['The Goblin']),
 ('Daniel Vízek', ['Ice Mother']),
 ('CC King', []),
 ('Sage Correa', ['Kidnap']),
 ('Brooke Norbury', ['Double Date']),
 ('Jacob Tremblay', ['Wonder', 'The Book of Henry']),
 ('Mckenna Grace', ['Gifted', 'I, Tonya']),
 ('Zhi-Xuan Wu', []),
 ('Balim Gaye Bayrak', []),
 ('Timothy Radford', ["The Zookeeper's Wife"]),
 ('Francisca Aronsson', ['Cebiche de tiburón']),
 ('Kairi Danielson', ['Celaderion'])]
```

<!-- grade_id: cell-dd0fbeed69b0c018, points: 2, grade: true, solution: false -->
```python
assert len(Q8()) == 20
assert all(len(row) == 2 for row in Q8())
```

__9. Select a sorted (ascending) and distinct list of movie genres__
* _hint_: Python provides a `set` structure and `sorted` function
* _return_: List[Genre] (where Genre = str)

<!-- grade_id: cell-e11e8ea4e390320a, points: 0, grade: false, solution: true -->
```python
def Q9():
    # YOUR CODE HERE
    # raise NotImplementedError()
    
    the_set = set()
    for i in basics.keys():
        the_set = the_set | set(basics[i]['genres'])
        
    return sorted(the_set)

Q9()
```

Output:
```
['Action',
 'Adult',
 'Adventure',
 'Animation',
 'Biography',
 'Comedy',
 'Crime',
 'Documentary',
 'Drama',
 'Family',
 'Fantasy',
 'Game-Show',
 'History',
 'Horror',
 'Music',
 'Musical',
 'Mystery',
 'News',
 'Reality-TV',
 'Romance',
 'Sci-Fi',
 'Short',
 'Sport',
 'Talk-Show',
 'Thriller',
 'War',
 'Western']
```

<!-- grade_id: cell-5c6cf9a4dd9a067a, points: 2, grade: true, solution: false -->
```python
assert len(Q9()) == 27
assert all(isinstance(x, str) for x in Q9())
```

__10. Select number of distinct movies for which a person is known for and which exists in the `basics` collection__
- _hint_: I insist on the distinct and exists word in the question
- _return_: Count (where Count = int)

<!-- grade_id: cell-2251a150ede884cd, points: 0, grade: false, solution: true -->
```python
def Q10():
    # YOUR CODE HERE
    # raise NotImplementedError()
    temp = []
    for i in names.find():
        temp.extend(i['knownForTitles'])
    temp.sort()
    set1 = set(temp)
    
    set2 = set(basics.keys())
    
    # print(set2 & set1)
    
    return len(set2 & set1)
        

Q10()
```

Output:
```
3393
```

<!-- grade_id: cell-eee5133254e9654b, points: 2, grade: true, solution: false -->
```python
assert isinstance(Q10(), int)
```