# Tutorial 6: The Value of Values

__The goal of this assignment is to answer 10 questions related to [a talk on programming from Rich Hickey](https://youtu.be/-6BsiVyC1kM).__

I think this talk is important to develop your critical regard on the current data management landscape.

I also hope this talk will open your mind on alternative solutions to store data and program software.

**For each question, your function should return the number associated to the answer.**

There is only one valid answer per question.

**Example: What is the sound of a dog ?**
1. meow
2. **woof**
3. hoot
4. honk

```python
def Q0():
    return 2
```

__Grade scale__: 10 points
- __correct answer__: 1 points
- __incorrect answer__: 0 points

```python
from IPython.display import HTML

HTML('<iframe width="980" height="600" src="https://www.youtube.com/embed/-6BsiVyC1kM" frameborder="0" allowfullscreen></iframe>')
```

Output:
```
<IPython.core.display.HTML object>
```

# Questions

__1. How would you define the concept of "Place Oriented Programming" ?__
1. Maintaining information in sorted order
2. Allocating memory prior to data processing
3. Replacing old information by newer information
4. Allowing programmer to travel around the world

<!-- grade_id: cell-af420023b7676ec1, points: 0, grade: false, solution: true -->
```python
def Q1():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 3
Q1()
```

Output:
```
3
```

<!-- grade_id: cell-38b42e490724ca6d, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q1(), int)
```

__2. What was the main reason to justify the early use of Place Oriented Programming ?__
1. It was the only way to code on a computer
2. Programmers at the time were all idiots
3. Memory resources were too limited
4. Human resources were too limited

<!-- grade_id: cell-b969bb555ff6e41f, points: 0, grade: false, solution: true -->
```python
def Q2():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 3
Q2()
```

Output:
```
3
```

<!-- grade_id: cell-12c7bfa02b131299, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q2(), int)
```

__3. A value is not ...__
1. shareable 
2. immutable
3. changeable
4. semantically transparent

<!-- grade_id: cell-4a5f68774ae2d9f4, points: 0, grade: false, solution: true -->
```python
def Q3():
    # YOUR CODE HERE
    #raise NotImplementedError()
    return 3
Q3()
```

Output:
```
3
```

<!-- grade_id: cell-d267f272cd2eca06, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q3(), int)
```

__4. Which data structure is not a value in Python ?__
1. an int
2. a tuple
3. a string
4. a list

<!-- grade_id: cell-5151e01f60804944, points: 0, grade: false, solution: true -->
```python
def Q4():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 4

Q4()
```

Output:
```
4
```

<!-- grade_id: cell-481746201e851626, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q4(), int)
```

__5. Which programming method should be forbidden on a value ?__
1. copy
2. save
3. accrete
4. overwrite

<!-- grade_id: cell-0d6b6981c043eae0, points: 0, grade: false, solution: true -->
```python
def Q5():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 4
Q5()
```

Output:
```
4
```

<!-- grade_id: cell-d7fc18b7ec8a407d, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q5(), int)
```

__6. Which important notion is incorporated to the notion of fact compared to a value ?__
1. location
2. space
3. time
4. change

<!-- grade_id: cell-ed5f0b03a1675a58, points: 0, grade: false, solution: true -->
```python
def Q6():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 3
Q6()
```

Output:
```
3
```

<!-- grade_id: cell-93ef4ae7112f9795, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q6(), int)
```

__7. What does a fact oriented system support ?__
1. decision making
2. improved disk usage
3. improved performance
4. focus on present time

<!-- grade_id: cell-6a1daef61811752d, points: 0, grade: false, solution: true -->
```python
def Q7():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 1
Q7()
```

Output:
```
1
```

<!-- grade_id: cell-fe9ea22acb67684f, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q7(), int)
```

__8. Which system already supports facts oriented programming ?__
1. source control
2. file systems
3. sql databases
4. nosql databases

<!-- grade_id: cell-803731b2fe93baa1, points: 0, grade: false, solution: true -->
```python
def Q8():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 1
Q8()
```

Output:
```
1
```

<!-- grade_id: cell-3916000bef44802c, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q8(), int)
```

__9. What is the main problem of Big Data according to Rich Hickey ?__
1. current systems only remember the last thing
2. we can't handle high volume of data
3. databases are not fast enough
4. MapReduce is too complex

<!-- grade_id: cell-c51ba59356d3f7f0, points: 0, grade: false, solution: true -->
```python
def Q9():
    # YOUR CODE HERE
    # raise NotImplementedError()
    return 4

Q9()
```

Output:
```
4
```

<!-- grade_id: cell-7a48f41b09b209a3, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q9(), int)
```

__10. Why do we still rely on place orientation for our programming languages and databases systems ?__
1. we lack computer resources
2. we cannot log our system information
3. the software community is not familiar with these concepts
4. the cost for setting an alternative system would be too high

<!-- grade_id: cell-c5c2bdae3fd3601f, points: 0, grade: false, solution: true -->
```python
def Q10():
    # YOUR CODE HERE
    #raise NotImplementedError()
    return 3
Q10()
```

Output:
```
3
```

<!-- grade_id: cell-3a108cbb112130ff, points: 1, grade: true, solution: false -->
```python
assert isinstance(Q10(), int)
```